## Εργασία 2 ##

Καλωσήρθατε στην δεύτερη σας εργασία. Η εργασία αυτή έχει σκοπό να σας βοηθήσει να εμπεδώσετε τα δενδρικά μοντέλα και την αξιολόγηση μοντέλων.

Στην εργασία αυτή θα πρέπει να συμπληρώσετε κώδικα Python 3 στα σημεία που αναφέρουν # YOUR CODE HERE. Μην τροποποιείτε τον κώδικα που βρίσκεται εκτός αυτών των περιοχών.

Πρωτού παραδόσετε την εργασία σας σιγουρευτείτε ότι ο κώδικας σε όλα τα κελιά τρέχει σωστά. Για το σκοπό αυτό επιλέξτε από το μενού Χρόνος εκτέλεσης (runtime) -> Επανεκίνηση περιόδου λειτουργίας και εκτέλεση όλων.

Συμπληρώστε το όνομα (NAME) και το AEM σας παρακάτω:

```python
NAME = "DIMITRA GALANI"
AEM = "10331"
```

**1** Διαβάστε το διαθέσιμο από την sklearn σύνολο δεδομένων diabetes που αφορά την πρόβλεψη της πορείας που θα έχει ένας ασθενής με διαβήτη ένα χρόνο μετά (εργασία παλινδρόμησης) και αποθηκεύστε το σε μια μεταβλητή με το όνομα diabetes. (0.5 μονάδες)

<!-- grade_id: cell-ac5d304e50a110a6, points: 0, grade: false, solution: true -->
```python
from sklearn.datasets import load_diabetes
diabetes = load_diabetes()
```

<!-- grade_id: cell-2d158133c851eb4e, points: 0.5, grade: true, solution: false -->
```python
"""Τεστ ορθής ανάγνωσης του συνόλου δεδομένων"""
assert diabetes.feature_names == ['age', 'sex', 'bmi', 'bp', 's1', 's2', 's3', 's4', 's5', 's6']
```

**2** Χωρίστε το σύνολο που διαβάσατε σε υποσύνολο ελέγχου 20% και υποσύνολο εκπαίδευσης 80% με την μέθοδο train_test_split και σπόρο τυχαιότητας 42. Aποθηκεύστε το σύνολο εκπαίδευσης σε μεταβλητές X_train, y_train και το σύνολο ελέγχου σε μεταβλητές X_test, y_test. (0.5 μονάδες)

<!-- grade_id: cell-a637aedb66ae6d89, points: 0, grade: false, solution: true -->
```python
from sklearn.model_selection import train_test_split
# Split the data into 0.2 test and 0.8 train
X_train, X_test, y_train, y_test = train_test_split(diabetes.data, diabetes.target, test_size = 0.2, random_state=42)
```

<!-- grade_id: cell-300331f01f617c3f, points: 0.5, grade: true, solution: false -->
```python
"""Τεστ ορθού διαχωρισμού του συνόλου δεδομένων"""
assert round(X_train[0][8], 5) == 0.02736
assert round(X_test[0][8], 5) == 0.03243
```

Output:
```
AssertionError                            Traceback (most recent call last)
C:\Users\Dimitra\AppData\Local\Temp\ipykernel_29248\4192753042.py in <module>
      1 """Τεστ ορθού διαχωρισμού του συνόλου δεδομένων"""
----> 2 assert round(X_train[0][8], 5) == 0.02736
      3 assert round(X_test[0][8], 5) == 0.03243

AssertionError:
```

**3** Χρησιμοποιώντας τα δεδομένα εκπαίδευσης, εκπαιδεύστε δύο δένδρα παλινδρόμησης με κριτήρια διαχωρισμού τα squared_eror και poisson ([δείτε ενότητα 1.10.7.2](https://scikit-learn.org/stable/modules/tree.html#tree)) αντίστοιχα, θέτοντας τον σπόρο τυχαιότητας στην τιμή 42 και το μέγιστο βάθος στην τιμή 3. Υπολογίστε τη μετρική r2 με βάση τις προβλέψεις τους στα δεδομένα ελέγχου, τις οποίες αποθηκεύστε στις μεταβλητές y_pred1 (squared_error) και y_pred2 (poisson) αντίστοιχα.
Αποθηκεύστε τις μετρικές στις μεταβλητές r2_squared_error και r2_poisson αντίστοιχα. (1 μονάδα)

<!-- grade_id: cell-62cfb517c0f9271e, points: 0, grade: false, solution: true -->
```python
from sklearn.tree import DecisionTreeRegressor
from sklearn.metrics import r2_score 

# Squared Error
squared_error = DecisionTreeRegressor(criterion = 'squared_error', random_state = 42, max_depth = 3) # Create the model with 1) split criteria mean squared error 2) random state = 42 3) max depth = 3
squared_error.fit(X_train, y_train) # Fit the model
y_pred1 = squared_error.predict(X_test)
r2_squared_error = r2_score(y_test, y_pred1)

# Poisson
poisson = DecisionTreeRegressor(criterion = 'poisson', random_state = 42, max_depth = 3) # Create the model with 1) split criteria poisson 2) random state = 42 3) max depth = 3
poisson.fit(X_train, y_train) # Fit the model
y_pred2 = poisson.predict(X_test)
r2_poisson = r2_score(y_test, y_pred2)
```

<!-- grade_id: cell-714ce0b2b256307e, points: 1, grade: true, solution: false -->
```python
"""Τεστ ορθού υπολογισμού των μετρικών"""
assert round(r2_squared_error, 3) == 0.329
assert round(r2_poisson, 3) == 0.391
assert len(y_pred1) == len(y_pred2)
```

**4** Αναπτύξτε κώδικα με βάση την ιδιότητα [tree_](https://scikit-learn.org/stable/auto_examples/tree/plot_unveil_tree_structure.html#sphx-glr-auto-examples-tree-plot-unveil-tree-structure-py) του παρακάτω μοντέλου, προκειμένου να εμφανίσετε το μονοπάτι που ακολουθήθηκε στο δένδρο για την πρόβλεψη της πρώτης περίπτωσης του συνόλου ελέγχου. Σε περίπτωση που ελέγχεται δεύτερη φορά μια μεταβλητή στο μονοπάτι, θα πρέπει να διατηρούμε τον τελευταίο και στενότερο περιορισμό. Ακολουθούν παραδείγματα ορθής λειτουργίας, στα οποία φαίνεται και η μορφή αν-τότε που πρέπει να έχει το μονοπάτι. (4 μονάδες)

```python
from sklearn.tree import export_text
model = DecisionTreeRegressor(max_depth=3, random_state=42)
model.fit(X_train, y_train)
```

Output:
```
DecisionTreeRegressor(max_depth=3, random_state=42)
```

<!-- grade_id: cell-1697542afe8f99a2, points: 0, grade: false, solution: true -->
```python
def get_path(tree, instance, feature_names):
    feature = tree.tree_.feature
    threshold = tree.tree_.threshold
    value = tree.tree_.value

    # Initialize variables
    current_node = 0
    path_conditions = {}

    # Traverse the tree to find the path for the instance
    while True:
        if feature[current_node] != -2:  # If not a leaf node
            # Get the feature name and threshold for the current node
            feature_index = feature[current_node]
            feature_name = feature_names[feature_index]
            threshold_value = threshold[current_node]

            # Determine the direction to move in the tree
            if instance[feature_index] <= threshold_value:
                direction = "<="
            else:
                direction = ">"

            # Add or update the condition for the current feature
            if feature_index in path_conditions:
                # If the feature has been encountered before
                prev_feature_name, prev_direction, prev_threshold_value = path_conditions[feature_index]
                if prev_direction == direction:
                    # If the directions are the same, update the condition only if the new threshold is stricter
                    if direction == "<=":
                        if threshold_value < prev_threshold_value:
                            path_conditions[feature_index] = (feature_name, direction, threshold_value)
                    else:
                        if threshold_value > prev_threshold_value:
                            path_conditions[feature_index] = (feature_name, direction, threshold_value)
                else:
                    # If the directions are different, keep both conditions
                    path_conditions[feature_index + 0.5] = (feature_name, direction, threshold_value)
                    path_conditions[feature_index] = (prev_feature_name, prev_direction, prev_threshold_value)
            else:
                # If it's the first encounter of the feature, add it to the conditions
                path_conditions[feature_index] = (feature_name, direction, threshold_value)

            # Move to the left or right child node
            if direction == "<=":
                current_node = tree.tree_.children_left[current_node]
            else:
                current_node = tree.tree_.children_right[current_node]
        else:
            # Reached a leaf node, add the predicted value and exit the loop
            predicted_value = value[current_node][0][0]
            if predicted_value == int(predicted_value):
                formatted_predicted = f"{int(predicted_value)}"
            elif int(predicted_value * 10) == predicted_value * 10:
                formatted_predicted = f"{predicted_value:.1f}"
            else:
                formatted_predicted = f"{predicted_value:.2f}"
            
            # Construct the path string using the last and stricter condition for each feature
            path_conditions_str = [f"{name} {direction} {value:.2f}" for _, (name, direction, value) in path_conditions.items()]
            path_conditions_str.append(f"τότε {formatted_predicted}")
            break
        
    return "αν " + " και ".join(path_conditions_str[:-1]) + " " + path_conditions_str[-1]

print(get_path(model, X_test[0], diabetes.feature_names))
print(get_path(model, X_test[1], diabetes.feature_names))
print(get_path(model, X_test[6], diabetes.feature_names))
```

Output:
```
αν bmi <= 0.01 και s5 > 0.01 και s4 <= 0.09 τότε 159.57
αν bmi > 0.01 και bmi <= 0.07 και s6 <= 0.03 τότε 175.8
αν bmi > 0.07 και s2 > 0.02 τότε 225.75
```

<!-- grade_id: cell-8aabb13a60a743dd, points: 4, grade: true, solution: false -->
```python
"""Τεστ ορθού υπολογισμού του μονοπατιού"""
assert get_path(model, X_test[0], diabetes.feature_names) == 'αν bmi <= 0.01 και s5 > 0.01 και s4 <= 0.09 τότε 159.57'
assert get_path(model, X_test[1], diabetes.feature_names) == 'αν bmi > 0.01 και bmi <= 0.07 και s6 <= 0.03 τότε 175.8'
assert get_path(model, X_test[6], diabetes.feature_names) == 'αν bmi > 0.07 και s2 > 0.02 τότε 225.75'
```

**5** Διαχωρίστε τα δεδομένα diabetes σε υποσύνολα εκπαίδευσης 60% (X_train, y_train), επικύρωσης 15% (X_val, y_val) και ελέγχου 25% (X_test, y_test) με σπόρο τυχαιότητας 42. (1 μονάδα)

<!-- grade_id: cell-c2448c38ed5250e1, points: 0, grade: false, solution: true -->
```python
# Split the data into 0.25 test and 0.75 temporary
X_temp, X_test, y_temp, y_test = train_test_split(diabetes.data, diabetes.target, test_size=0.25, random_state=42)

# Split the temporary data into 0.75 * 0.2 = 0.15 validation and 0.75 * 0.8 = 0.6 train
X_train, X_val, y_train, y_val = train_test_split(X_temp, y_temp, test_size=0.2, random_state=42)

```

<!-- grade_id: cell-8da00aed9ea7c071, points: 1, grade: true, solution: false -->
```python
"""Τεστ ορθού διαχωρισμού του συνόλου δεδομένων"""
assert round(X_train[0][8], 5) == -0.08024
assert round(X_val[0][8], 5) == 0.09924
assert round(X_test[0][8], 5) == 0.03243
```

**6** Υλοποιήστε συνάρτηση, η οποία δοθέντος ενός δενδρικού μοντέλου παλινδρόμησης, διερευνά όλες τις ακέραιες τιμές της μεταβλητής min_samples_leaf από ένα κάτω όριο μέχρι ένα άνω όριο και επιστρέφει την βέλτιστη με βάση το mean squared error στο παραπάνω σύνολο επικύρωσης έπειτα από εκπαίδευση στο παραπάνω σύνολο εκπαίδευσης. Επιπλέον επιστρέφει και το mean squared error (στρογγυλοποιημένο σε ακέραιο) που πετυχαίνει στο σύνολο ελέγχου το μοντέλο με την βέλτιστη επίδοση στο σύνολο επικύρωσης. (2 μονάδες)

<!-- grade_id: cell-45e2ac4f9c5115f9, points: 0, grade: false, solution: true -->
```python
from sklearn.metrics import mean_squared_error

def best_min_samples_leaf(model, min, max):
    best_leaf = None
    best_mse_val = float('inf')
    best_mse_test = None

    for min_samples in range(min, max + 1):
        # Set the min_samples_leaf parameter
        model.set_params(min_samples_leaf=min_samples)
        
        # Train the model on the training set
        model.fit(X_train, y_train)
        
        # Predict on the validation set
        y_val_pred = model.predict(X_val)
        
        # Calculate MSE on the validation set
        mse_val = mean_squared_error(y_val, y_val_pred)
        
        # Update the best parameters if the current MSE is lower
        if mse_val < best_mse_val:
            best_mse_val = mse_val
            best_leaf = min_samples
            
            # Predict on the test set with the current best model
            y_test_pred = model.predict(X_test)
            best_mse_test = mean_squared_error(y_test, y_test_pred)
    
    # Return the best min_samples_leaf and the rounded MSE on the test set
    return best_leaf, int(round(best_mse_test))
```

<!-- grade_id: cell-0d6be21e54be9665, points: 2, grade: true, solution: false -->
```python
"""Τεστ ορθής υλοποίησης συνάρτησης"""
model = DecisionTreeRegressor(max_depth=3, random_state=42)
assert best_min_samples_leaf(model, 1, 8) == (7, 3358)
assert best_min_samples_leaf(model, 5, 15) == (11, 3271)
```

**7** Υλοποιήστε συνάρτηση η οποία δοθέντος ενός μοντέλου παλινδρόμησης, ενός συνόλου δεδομένων και μίας μετρικής, το αξιολογεί με τη μέθοδο της σταυρωτής επικύρωσης "άφησε ένα εκτός". (1 μονάδα)

<!-- grade_id: cell-880c64392f85bccb, points: 0, grade: false, solution: true -->
```python
from sklearn.model_selection import LeaveOneOut
import numpy as np

def leave_one_out(model, X, y, metric):
    # Initialize LeaveOneOut object
    loo = LeaveOneOut()
    scores = []

    # Loop over each train-test split
    for train_index, test_index in loo.split(X):
        # Split data
        X_train, X_test = X[train_index], X[test_index]
        y_train, y_test = y[train_index], y[test_index]

        # Train the model
        model.fit(X_train, y_train)
        
        # Evaluate the model
        y_pred = model.predict(X_test)
        
        # Compute the specified metric and append to list
        score = metric(y_test, y_pred)
        scores.append(score)

    # Compute the average of the scores
    average_score = np.mean(scores)
    return average_score
```

<!-- grade_id: cell-4afd7b6de1cd1f7f, points: 1, grade: true, solution: false -->
```python
"""Τεστ ορθής υλοποίησης συνάρτησης"""
model = DecisionTreeRegressor(max_leaf_nodes=6, min_samples_leaf=2, random_state=42)
assert round(leave_one_out(model, diabetes.data, diabetes.target, mean_squared_error), 0) == 3354
```